// Numbas version: finer_feedback_settings
{"name": "Test_deg_selv_tallregning", "metadata": {"description": "", "licence": "None specified"}, "duration": 0, "percentPass": 0, "showQuestionGroupNames": false, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Group", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": [""], "variable_overrides": [[]], "questions": [{"name": "Tall_testdegselv", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "", "advice": "<p>Husk p&aring; regnerekkef&oslash;lgen:</p>\n<ol>\n<li>Parenteser</li>\n<li>Potenser</li>\n<li>Multiplikasjon og Divisjon</li>\n<li>Addisjon og Subtraksjon</li>\n</ol>\n<p>I tillegg m&aring; vi huske &aring; regne fra&nbsp;<strong>venstre til h&oslash;yre</strong>, spesielt n&aring;r vi har flere operasjoner som gjennomf&oslash;res sammtidig!</p>\n<div data-jme-visible=\"cond_neg\">\n<p><span data-jme-visible=\"cond_neg23\">Dersom vi skal plusse p&aring; noe som er negativt, er det samme som &aring; subtrahere med det samme tallet!</span><br/><span data-jme-visible=\"cond_neg23\">Husk:</span></p>\n</div>\n<ul>\n<li><span data-jme-visible=\"cond_neg23\">++ blir +</span></li>\n<li><span data-jme-visible=\"cond_neg23\">--&nbsp; &nbsp;blir +</span></li>\n<li><span data-jme-visible=\"cond_neg23\">+-&nbsp; blir -</span></li>\n<li><span data-jme-visible=\"cond_neg23\">-+&nbsp; blir -</span><span data-jme-visible=\"cond_neg23\"></span></li>\n</ul>\n<p><span data-jme-visible=\"cond1\">Husk at $a^1=a$ (multiplisert med seg selv $1$ gang)</span></p>\n<p><span data-jme-visible=\"cond0\">Husk at $a^0=1$ (multiplisert med seg selv $0$ ganger)&nbsp;</span></p>", "rulesets": {}, "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"d22": {"name": "d22", "group": "Tall 2.2", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b23": {"name": "b23", "group": "Tall 2.3", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c23": {"name": "c23", "group": "Tall 2.3", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "d31": {"name": "d31", "group": "Tall 3.1", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "f31": {"name": "f31", "group": "Tall 3.1", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "a32": {"name": "a32", "group": "Tall 3.2", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "a31": {"name": "a31", "group": "Tall 3.1", "definition": "random(0..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b31": {"name": "b31", "group": "Tall 3.1", "definition": "if(a31<=3,random(0..3),random(0,2))", "description": "", "templateType": "anything", "can_override": false}, "c22": {"name": "c22", "group": "Tall 2.2", "definition": "random(-5..5)", "description": "", "templateType": "anything", "can_override": false}, "b22": {"name": "b22", "group": "Tall 2.2", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c32": {"name": "c32", "group": "Tall 3.2", "definition": "random(0..3)", "description": "", "templateType": "anything", "can_override": false}, "f32": {"name": "f32", "group": "Tall 3.2", "definition": "random(-5..5 except 0 except 1 except -1)", "description": "", "templateType": "anything", "can_override": false}, "d32": {"name": "d32", "group": "Tall 3.2", "definition": "random(-2..2 except 0 except 1)*f32", "description": "", "templateType": "anything", "can_override": false}, "e32": {"name": "e32", "group": "Tall 3.2", "definition": "random(0..3 except 0)", "description": "", "templateType": "anything", "can_override": false}, "a22": {"name": "a22", "group": "Tall 2.2", "definition": "random(-5..5 except 0)*b22", "description": "", "templateType": "anything", "can_override": false}, "e22": {"name": "e22", "group": "Tall 2.2", "definition": "random(-5..5)", "description": "", "templateType": "anything", "can_override": false}, "a21": {"name": "a21", "group": "2.1", "definition": "random(0..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b21": {"name": "b21", "group": "2.1", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c21": {"name": "c21", "group": "2.1", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "d21": {"name": "d21", "group": "2.1", "definition": "random(0..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "e21": {"name": "e21", "group": "2.1", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "f21": {"name": "f21", "group": "2.1", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "a12": {"name": "a12", "group": "Tall 1.2", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b12": {"name": "b12", "group": "Tall 1.2", "definition": "random(4..7)", "description": "", "templateType": "anything", "can_override": false}, "c12": {"name": "c12", "group": "Tall 1.2", "definition": "random(4..10)", "description": "", "templateType": "anything", "can_override": false}, "a11": {"name": "a11", "group": "Tall 1.1", "definition": "random(0..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b11": {"name": "b11", "group": "Tall 1.1", "definition": "random(0..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c11": {"name": "c11", "group": "Tall 1.1", "definition": "random(0..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "c31": {"name": "c31", "group": "Tall 3.1", "definition": "random(-5..5)", "description": "", "templateType": "anything", "can_override": false}, "e31": {"name": "e31", "group": "Tall 3.1", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "b32": {"name": "b32", "group": "Tall 3.2", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "a23": {"name": "a23", "group": "Tall 2.3", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "d23": {"name": "d23", "group": "Tall 2.3", "definition": "random(-5..5 except 0)", "description": "", "templateType": "anything", "can_override": false}, "e23": {"name": "e23", "group": "Tall 2.3", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "f23": {"name": "f23", "group": "Tall 2.3", "definition": "random(-10..10 except 0)", "description": "", "templateType": "anything", "can_override": false}, "cond_neg": {"name": "cond_neg", "group": "Conditions", "definition": "b22<0||d22<0||e22<0", "description": "", "templateType": "anything", "can_override": false}, "cond_neg31": {"name": "cond_neg31", "group": "Conditions", "definition": "c31<0||e31<0||f31<0", "description": "", "templateType": "anything", "can_override": false}, "cond_neg3": {"name": "cond_neg3", "group": "Conditions", "definition": "d32<0||f32<0", "description": "", "templateType": "anything", "can_override": false}, "cond_neg23": {"name": "cond_neg23", "group": "Conditions", "definition": "c23<0||d23<0||e23<0", "description": "", "templateType": "anything", "can_override": false}, "cond0": {"name": "cond0", "group": "Conditions", "definition": "b31=0||c32=0||e32=0", "description": "", "templateType": "anything", "can_override": false}, "cond1": {"name": "cond1", "group": "Conditions", "definition": "b31=1||c32=1||e32=1", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "2.1", "variables": ["a21", "b21", "c21", "d21", "e21", "f21"]}, {"name": "Tall 1.2", "variables": ["a12", "b12", "c12"]}, {"name": "Tall 1.1", "variables": ["a11", "b11", "c11"]}, {"name": "Tall 2.2", "variables": ["a22", "b22", "c22", "d22", "e22"]}, {"name": "Tall 3.1", "variables": ["a31", "b31", "c31", "d31", "e31", "f31"]}, {"name": "Tall 3.2", "variables": ["a32", "b32", "c32", "d32", "e32", "f32"]}, {"name": "Tall 2.3", "variables": ["a23", "b23", "c23", "d23", "e23", "f23"]}, {"name": "Conditions", "variables": ["cond_neg", "cond_neg31", "cond_neg3", "cond_neg23", "cond0", "cond1"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "F\u00f8rste sp\u00f8rsm\u00e5l tall", "marks": "56", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Tall 3.1", "rawLabel": "Tall 3.1", "otherPart": 2, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Tall 1.2", "rawLabel": "Tall 1.2", "otherPart": 5, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p>Regn ut&nbsp; <!--StartFragment --><span class=\"cf0\">$\\var{a22} \\div$ <span data-jme-visible=\"b22&lt;0\">$($ </span>$\\var{b22}$ <span data-jme-visible=\"b22&lt;0\">$)$ </span>$-$ <span data-jme-visible=\"c22&lt;0\">$($</span>$\\var{c22}$ <span data-jme-visible=\"c22&lt;0\">$)$</span>$+$ <span data-jme-visible=\"d22&lt;0\">$($</span>$\\var{d22}$<span data-jme-visible=\"d22&lt;0\">$)$</span> $\\cdot$</span><span class=\"cf0\">&nbsp;<span data-jme-visible=\"e22&lt;0\">$($</span>$\\var{e22}$<span data-jme-visible=\"e22&lt;0\">$)$</span></span><!--EndFragment --></p>", "minValue": "a22/b22-c22+d22*e22", "maxValue": "a22/b22-c22+d22*e22", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 2.3", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p style=\"padding-left: 40px;\">Regn ut $\\var{a23} \\cdot\\Bigl( \\var{b23}-$<span data-jme-visible=\"c23&lt;0\">$($</span>$\\var{c23}$<span data-jme-visible=\"c23&lt;0\">$)$</span>$\\Bigr)+$<span data-jme-visible=\"d23&lt;0\">$($</span>$\\var{d23}$<span data-jme-visible=\"d23&lt;0\">$)$</span>$\\cdot \\Bigl(\\var{e23}-$<span data-jme-visible=\"f23&lt;0\">$($</span>$\\var{f23}$<span data-jme-visible=\"f23&lt;0\">$)$</span>$\\Bigr)$</p>", "minValue": "a23*(b23-c23)+d23*(e23-f23)", "maxValue": "a23*(b23-c23)+d23*(e23-f23)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 3.1", "marks": "29", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Tall 3.2", "rawLabel": "Tall 3.2", "otherPart": 3, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Tall 2.3", "rawLabel": "Tall 2.3", "otherPart": 1, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p style=\"padding-left: 40px;\">Regn ut&nbsp; $\\var{a31}^\\var{b31}-$<span data-jme-visible=\"c31&lt;0\">$($</span>$\\var{c31}$<span data-jme-visible=\"c31&lt;0\">$)$</span>$\\cdot\\Bigl(\\var{d31}+$<span data-jme-visible=\"e31&lt;0\">$($</span>$\\var{e31}$<span data-jme-visible=\"e31&lt;0\">$)$</span>$\\Bigr)+$<span data-jme-visible=\"f31&lt;0\">$($</span>$\\var{f31}$<span data-jme-visible=\"f31&lt;0\">$)$</span></p>", "minValue": "a31^b31-c31*(d31+e31)+f31", "maxValue": "a31^b31-c31*(d31+e31)+f31", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 3.2", "marks": "15", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p>Regn ut $\\var{a32}\\cdot\\Bigl( $<span data-jme-visible=\"b32&lt;0\">$($</span>$\\var{b32}$<span data-jme-visible=\"b32&lt;0\">$)$</span>$^\\var{c32}-$<span data-jme-visible=\"d32&lt;0\">$($</span>$\\var{d32}$<span data-jme-visible=\"d32&lt;0\">$)$</span>$^\\var{e32}\\div$<span data-jme-visible=\"f32&lt;0\">$($</span>$\\var{f32}$<span data-jme-visible=\"f32&lt;0\">$)$</span>$ \\Bigr)$</p>", "minValue": "a32*(b32^c32-(d32^e32/f32))", "maxValue": "a32*(b32^c32-(d32^e32/f32))", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 2.1", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p>Regn ut&nbsp;</p>\n<p>$\\var{a21} \\cdot\\Bigl( \\var{b21}-$<span data-jme-visible=\"c21&lt;0\">$($</span>$\\var{c21}$<span data-jme-visible=\"c21&lt;0\">$)$</span>$\\Bigr)+$<span data-jme-visible=\"d21&lt;0\">$($</span>$\\var{d21}$<span data-jme-visible=\"d21&lt;0\">$)$</span>$\\cdot \\Bigl(\\var{e21}-$<span data-jme-visible=\"f21&lt;0\">$($</span>$\\var{f21}$<span data-jme-visible=\"f21&lt;0\">$)$</span>$\\Bigr)$</p>", "minValue": "a21*(b21-c21)+d21*(e21-f21)", "maxValue": "a21*(b21-c21)+d21*(e21-f21)", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 1.2", "marks": "28", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Tall 2.1", "rawLabel": "Tall 2.1", "otherPart": 4, "variableReplacements": [], "availabilityCondition": "answered and credit=1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}, {"label": "Tall 1.1", "rawLabel": "Tall 1.1", "otherPart": 6, "variableReplacements": [], "availabilityCondition": "answered and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": true}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p>Regn ut&nbsp;$ \\var{a12} - \\var{b12}&nbsp;\\cdot \\var{c12}&nbsp;$</p>", "minValue": "a12-b12*c12", "maxValue": "a12-b12*c12", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "Tall 1.1", "marks": "14", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Tall", "prompt": "<p>Regn ut $\\var{a11}\\cdot \\var{b11}+\\var{c11}$</p>", "minValue": "a11*b11+c11", "maxValue": "a11*b11+c11", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "partsMode": "explore", "maxMarks": "100", "objectives": [{"name": "Tall", "limit": "100"}], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always"}]}], "allowPrinting": true, "navigation": {"allowregen": true, "reverse": true, "browse": true, "allowsteps": true, "showfrontpage": true, "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "typeendtoleave": false, "startpassword": "", "autoSubmit": true, "allowAttemptDownload": false, "downloadEncryptionKey": "", "showresultspage": "oncompletion"}, "timing": {"allowPause": true, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"enterreviewmodeimmediately": true, "showactualmarkwhen": "always", "showtotalmarkwhen": "always", "showanswerstatewhen": "always", "showpartfeedbackmessageswhen": "always", "showexpectedanswerswhen": "inreview", "showadvicewhen": "inreview", "allowrevealanswer": true, "intro": "", "end_message": "", "results_options": {"printquestions": true, "printadvice": true}, "feedbackmessages": [], "reviewshowexpectedanswer": true, "showanswerstate": true, "reviewshowfeedback": true, "showactualmark": true, "showtotalmark": true, "reviewshowscore": true, "reviewshowadvice": true}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "Francesca Azzolini", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/19466/"}], "extensions": [], "custom_part_types": [], "resources": []}